# Дополнительное задание

### Библиотеки
Во всех задачах можно использовать библиотеки `requests` и `BeautifulSoup` (`bs4`).

### Задача 1
Написать функцию `get_news(url)`, принимающую на вход адрес веб-страницы `url`, загружающую эту веб-страницу и проверяющую, встречается ли в ней слово `Harry` (с большой буквы). Функция должна возвращать `True`, если встречается, и `False` в противном случае. Также функция должна возвращать `False`, если страницу не удалось открыть (например, была получена ошибка *404 Not Found*.)

<!-- grade_id: pr_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from bs4 import BeautifulSoup        # Для обработки HTML
import re
import requests
def get_news(url):
    try:
        page = requests.get(url,timeout=5)
        cont = BeautifulSoup(page.content)
        
    except requests.exceptions.HTTPError:
        return False
    else:
        for i in cont.find_all():
            if ("Harry" in i.text):
                return True
        return False

print (get_news("https://en.wikipedia.org/w/index.php?title=J._K._Rowling&oldid=694008857"))
print (get_news("https://en.wikipedia.org/w/index.php?title=Star_Wars&oldid=694701430"))
print (get_news("https://en.wikipedia.org/w/index.php?title=Darth_Vader&oldid=694617684"))
```

Output:
```
True
True
False
```

<!-- grade_id: pr_1_1, points: 1, grade: true, solution: false -->
```python
assert get_news("https://en.wikipedia.org/w/index.php?title=J._K._Rowling&oldid=694008857")
assert get_news("https://en.wikipedia.org/w/index.php?title=Star_Wars&oldid=694701430")
assert not get_news("https://en.wikipedia.org/w/index.php?title=Darth_Vader&oldid=694617684")
```

<!-- grade_id: pr_1_2, points: 1, grade: true, solution: false -->
```python
print (get_news("http://math-info.hse.ru/there_is_no_Harry_here"))
```

Output:
```
False
```

### Задача 2
Написать функцию `get_strong(html)`, принимающую на вход html-страницу в виде длинной строки, записанной в переменную `html`, и возвращающую строчку, содержащуюся в первом теге `strong`.

Примеры см. в тестах.

**Подсказка.** Вы можете создать объект `BeautifulSoup`.

<!-- grade_id: pr_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from bs4 import BeautifulSoup
import re
def get_strong(html):
    cont = BeautifulSoup(html)
    for i in cont.find_all("strong"):
        return i.text
    
    
print (get_strong("<html><body><p>Hello, <strong>World</strong>!")) #"World"
html = """<html>
    <body>
        <p>
            Hello,
            <strong>
                World
            </strong>
        </p>
    </body>
</html>"""
    
print (get_strong(html).strip()) #"World"
print (get_strong("<html><body><p>tag &lt;strong&gt; is used in HTML\n to make letters <strong>stronger</strong>")) #"stronger"
print (get_strong("<html><body><strong>One\nTwo</strong></body></html>")) #"One\nTwo"
```

Output:
```
World
World
stronger
One
Two
```

<!-- grade_id: pr_2_2, points: 1, grade: true, solution: false -->
```python
assert get_strong("<html><body><p>Hello, <strong>World</strong>!") == "World"
html = """<html>
    <body>
        <p>
            Hello,
            <strong>
                World
            </strong>
        </p>
    </body>
</html>"""
    
assert get_strong(html).strip() == "World"
assert get_strong("<html><body><p>tag &lt;strong&gt; is used in HTML\n to make letters <strong>stronger</strong>") == "stronger"
assert get_strong("<html><body><strong>One\nTwo</strong></body></html>") == "One\nTwo"
```

### Задача 3
Для вставки картинок в HTML используется тег `<img>`, содержащий параметр `src` — адрес файла с картинкой. Например, `<img src="https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"/>`. Написать функцию `all_images_src(html)`, принимающую на вход длинную строчку с HTML-документом, а возвращающую список адресов всех картинок, встречающихся в этом документе (в том порядке, в котором они встречаются в документе).

<!-- grade_id: pr_3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from bs4 import BeautifulSoup
import re
def all_images_src(html):
    cont = BeautifulSoup(html)
    arr = cont.find_all("img", src = True)
    real_arr = []
    for i in arr:
        i = str(i)        
        i = re.search(r'src=(.+\.jpg)|(.+\.png)', i)     #https?://[^\"\s]+
        i = i.group(0)
        i=re.sub(".*=\"", "", i)
        real_arr.append(i)
    return real_arr


print(all_images_src('<html><body><img src="https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"/>'))
print (all_images_src('<html><body><IMG src="test.jpg">\n'
                        '<p>Some text\n'
                        '<img SRC=\'well.png\'>\n'
                        '</p></body></html>'))
print (all_images_src('<html><body><p><a href="link.html">'
                      '<img alt="Just a test image" src="this is a test.jpg"><ul>' + "\n"
                      .join("<li><img src='img%04i.png'></li>" % i for i in range(1000)) + 
                      "</ul></p></body></html>"))


```

Output:
```
['https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg']
['test.jpg', 'well.png']
['this is a test.jpg', 'img0000.png', 'img0001.png', 'img0002.png', 'img0003.png', 'img0004.png', 'img0005.png', 'img0006.png', 'img0007.png', 'img0008.png', 'img0009.png', 'img0010.png', 'img0011.png', 'img0012.png', 'img0013.png', 'img0014.png', 'img0015.png', 'img0016.png', 'img0017.png', 'img0018.png', 'img0019.png', 'img0020.png', 'img0021.png', 'img0022.png', 'img0023.png', 'img0024.png', 'img0025.png', 'img0026.png', 'img0027.png', 'img0028.png', 'img0029.png', 'img0030.png', 'img0031.png', 'img0032.png', 'img0033.png', 'img0034.png', 'img0035.png', 'img0036.png', 'img0037.png', 'img0038.png', 'img0039.png', 'img0040.png', 'img0041.png', 'img0042.png', 'img0043.png', 'img0044.png', 'img0045.png', 'img0046.png', 'img0047.png', 'img0048.png', 'img0049.png', 'img0050.png', 'img0051.png', 'img0052.png', 'img0053.png', 'img0054.png', 'img0055.png', 'img0056.png
```

<!-- grade_id: pr_3_t, points: 1, grade: true, solution: false -->
```python
assert all_images_src('<html><body><img src="https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"/>') == ['https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg']
assert all_images_src( ('<html><body><IMG src="test.jpg">\n'
                        '<p>Some text\n'
                        '<img SRC=\'well.png\'>\n'
                        '</p></body></html>') ) == ["test.jpg", "well.png"]
assert all_images_src('<html><body><p><a href="link.html">'
                      '<img alt="Just a test image" src="this is a test.jpg"><ul>' + "\n"
                      .join("<li><img src='img%04i.png'></li>" % i for i in range(1000)) + 
                      "</ul></p></body></html>"                     ) == ['this is a test.jpg'] + ['img%04i.png' % i for i in range(1000)]
```

<!-- grade_id: pr_4, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: pr_4_t, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: pr_5, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: pr_5_1, points: 4, grade: true, solution: false -->
```python

```

```python

```